Photoreceptor loss does not recruit neutrophils despite strong microglial activation

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #2 (Public review):

Summary:

This study uses in vivo multimodal high-resolution imaging to track how microglia and neutrophils respond to light-induced retinal injury from soon after injury to 2 months post-injury. The in vivo imaging finding was subsequently verified by ex vivo study. The results suggest that despite the highly active microglia at the injury site, neutrophils were not recruited in response to acute light-induced retinal injury.

Strengths:

An extremely thorough examination of the cellular-level immune activity at the injury site. In vivo imaging observations being verified using ex vivo techniques is a strong plus.

Thank you!

Weaknesses:

This paper is extremely long, and in the perspective of this reviewer, needs to be better organized. Update: Modifications have been made throughout, which has made the manuscript easier to follow.

Thank you!

Study weakness: though the finding prompts more questions and future studies, the findings discussed in this paper is potentially important for us to understand how the immune cells respond differently to different severity level of injury. The study also demonstrated an imaging technology which may help us better understand cellular activity in living tissue during earlier time points.

We agree that AOSLO has much to offer and this represents some of the earliest reports of its kind.

Comments on revisions:

I appreciate the thorough clarification and re-organization by the authors, and the messages in the manuscript are now more apparent. I recommend also briefly discussing limitations/future improvements in the discussion or conclusion.

We have added a section to the discussion entitled “Limitations and future improvements”, please see lines 665 – 677.

Reviewer #3 (Public review):

Summary

This work investigated the immune response in the murine retina after focal laser lesions. These lesions are made with close to 2 orders of magnitude lower laser power than the more prevalent choroidal neovascularization model of laser ablation. Histology and OCT together show that the laser insult is localized to the photoreceptors and spares the inner retina, the vasculature and the pigment epithelium. As early as 1-day after injury, a loss of cell bodies in the outer nuclear layer is observed. This is accompanied by strong microglial proliferation to the site of injury in the outer retina where microglia do not typically reside. The injury did not seem to result in the extravasation of neutrophils from the capillary network, constituting one of the main findings of the paper. The demonstrated paradigm of studying the immune response and potentially retinal remodeling in the future in vivo is valuable and would appeal to a broad audience in visual neuroscience.

Strengths

Adaptive optics imaging of murine retina is cutting edge and enables non-destructive visualization of fluorescently labeled cells in the milieu of retinal injury. As may be obvious, this in vivo approach is a benefit for studying fast and dynamic immune processes on a local time scale - minutes and hours, and also for the longer days-to-months follow-up of retinal remodeling as demonstrated in the article. In certain cases, the in vivo findings are corroborated with histology.

Thank you!

The analysis is sound and accompanied by stunning video and static imagery. A few different sets of mouse models are used, a) two different mouse lines, each with a fluorescent tag for neutrophils and microglia, b) two different models of inflammation - endotoxin-induced uveitis (EAU) and laser ablation are used to study differences in the immune interaction.

Thank you!

One of the major advances in this article is the development of the laser ablation model for 'mild' retinal damage as an alternative to the more severe neovascularization models. This model would potentially allow for controlling the size, depth and severity of the laser injury opening interesting avenues for future study.

Thank you!

The time-course, 2D and 3D spatial activation pattern of microglial activation are striking and provide an unprecedented view of the retinal response to mild injury.

We agree that this more complete spatial and temporal evaluation made possible by in vivo imaging is novel.

Weaknesses

Generalization of the (lack of) neutrophil response to photoreceptor loss - there is ample evidence in literature that neutrophils are heavily recruited in response to severe retinal damage that includes photoreceptor loss. Why the same was not observed here in this article remains an open question. One could hypothesize that neutrophil recruitment might indeed occur under conditions that are more in line with the more extreme damage models, for example, with a stronger and global ablation (substantially more photoreceptor loss over a larger area). This parameter space is unwieldy and sufficiently large to address the question conclusively in the current article, i.e. how much photoreceptor loss leads to neutrophil recruitment? By the same token, the strong and general conclusion in the title - Photoreceptor loss does not recruit neutrophils - cannot be made until an exhaustive exploration be made of the same parameter space. A scaling back may help here, to reflect the specific, mild form of laser damage explored here, for instance - Mild photoreceptor loss does not recruit neutrophils despite...

We are striving for clarity and accuracy in our title without adding too many qualifiers. At present, we feel that the title as submitted is consistent and aligned with the central finding of our manuscript. The nuance that the reviewer points to is elaborated in the body of the manuscript and we hope the general readership appreciates the same level of detail as appreciated by reviewer #3.

EIU model - The EIU model was used as a positive control for neutrophil extravasation. Prior work with flow cytometry has shown a substantial increase in neutrophil counts in the EIU model. Yet, in all, the entire article shows exactly 2 examples in vivo and 3 ex vivo (Figure 7) of extravasated neutrophils from the EIU model (n = 2 mice). The general conclusion made about neutrophil recruitment (or lack thereof) is built partly upon this positive control experiment. But these limited examples, especially in the case where literature reports a preponderance of extravasated neutrophils, raise a question on the paradigm(s) used to evaluate this effect in the mild laser damage model.

This is a helpful suggestion. We agree that readers should see more evidence of the positive control. Therefore we have now included two more supplementary files that show that there is a strong neutrophil response to EIU. In Figure 7 – supplementary figure 1, we show many Ly-6G-positive neutrophils in the retina seen with histology at the 24 hour time point. In Figure 7 – video 3, we show massive Catchup-positive neutrophil presence in vivo at 24hrs as well. This aligns with our positive control and also the literature.

Overall, the strengths outweigh the weaknesses, provided the conclusions/interpretations are reconsidered.

With the added clarification about the magnitude of the neutrophil response in EIU, we feel that the conclusions presented in the manuscript as-is are valid and appropriate.

Recommendations for the authors:

Reviewer #3 (Recommendations for the authors):

The authors are applauded for embracing the reviewers' feedback and making substantial revisions. Some minor comments below:

The weakness noted in the public review encourages the authors to reconsider the interpretations drawn based on the results. One would have expected to see far more examples of extravasated neutrophils from the EIU model. That this was not seen weakens the neutrophil recruitment claim substantially. Even without this claim, the methods, laser damage model, time-course and spatial activation pattern of microglial activation are all striking and unprecedented. So, as stated in the public review, the strengths do indeed outweigh the weaknesses once the neutrophil claim is softened.

We address this in the response above. A strong neutrophil response was observed to EIU. This was confirmed with both histology and in vivo imaging.

This was alluded to by Reviewer 1 in the prior review - at times, there is an overemphasis on imaging technology that distracts from the scientific questions. The imaging is undoubtedly cutting-edge but also documented in prior work by the authors. Any efforts to reduce or balance the emphasis would help with the general flow.

Given that these discoveries are made possible partly through new technology, we prefer to keep the details of the innovation in the current manuscript. Given the exceptionally large readership of eLife, we feel some description of the AOSLO imaging is warranted in the manuscript.

Changes from the earlier version to the revised version

Abstract

Removed in the revised version: Abstract

Earlier version: In response to central nervous system (CNS) injury, tissue resident immune cells such as microglia and circulating systemic neutrophils are often first responders. The degree to which these cells interact in response to CNS damage is poorly understood, and even less so, in the neural retina which poses a challenge for high resolution imaging in vivo. In this study, we deploy fluorescence adaptive optics scanning light ophthalmoscopy (AOSLO) to study fluorescent microglia and neutrophils in mice. We simultaneously track immune cell dynamics using label-free phase-contrast AOSLO at micron-level resolution. Retinal lesions were induced with 488 nm light focused onto photoreceptor (PR) outer segments. These lesions focally ablated PRs, with minimal collateral damage to cells above and below the plane of focus. We used in vivo (AOSLO, SLO and OCT) imaging to reveal the natural history of the microglial and neutrophil response from minutes-to-months after injury. While microglia showed dynamic and progressive immune response with cells migrating into the injury locus within 1-day after injury, neutrophils were not recruited despite close proximity to vessels carrying neutrophils only microns away. Post-mortem confocal microscopy confirmed in vivo findings. This work illustrates that microglial activation does not recruit neutrophils in response to acute, focal loss of PRs, a condition encountered in many retinal diseases.

Revised version: In response to central nervous system (CNS) injury, tissue resident immune cells such as microglia and circulating systemic neutrophils are often first responders. The degree to which these cells interact in response to CNS damage is poorly understood, and even less so, in the neural retina which poses a challenge for high resolution imaging in vivo. In this study, we deploy fluorescence adaptive optics scanning light ophthalmoscopy (AOSLO) to study microglia and neutrophils in mice. We simultaneously track immune cell dynamics using label-free phase-contrast AOSLO at micron-level resolution. Retinal lesions were induced with 488 nm light focused onto photoreceptor (PR) outer segments. These lesions focally ablated PRs, with minimal collateral damage to cells above and below the plane of focus. We used in vivo (AOSLO, SLO and OCT) imaging to reveal the natural history of the microglial and neutrophil response from minutes-to-months after injury. While microglia showed dynamic and progressive immune response with cells migrating into the injury locus within 1-day after injury, neutrophils were not recruited despite close proximity to vessels carrying neutrophils only microns away. Post-mortem confocal microscopy confirmed in vivo findings. This work illustrates that microglial activation does not recruit neutrophils in response to acute, focal loss of PRs, a condition encountered in many retinal diseases.

Introduction

In the mammalian retina, a rapid and coordinated immune response to infection or injury is important for maintaining tissue homeostasis. This is especially critical in the eye since mature retinal neurons do not typically regenerate, resulting in long-term functional losses for the host.1,2 The retina is considered immune-privileged and is equipped with a resident population of innate immune cells including microglia.3,4 In healthy retina, microglia are distributed primarily in the inner retina, residing within nerve fiber layer (NFL), inner plexiform layer (IPL) and outer plexiform layers (OPL).5,6, generally avoiding the nuclear layers. Microglia tile the retina and like their counterparts in the brain, exhibit long, thin processes that continually probe the neuro-glial microenvironment.4,7,8

In addition to phagocytosing debris9,10, regulating synaptic maintenance11,12 and removing dead tissue13,14, microglia can secrete chemokines to recruit other leukocytes to help fight infection and repair damaged tissue.4,15–17 For many injuries, one of the first systemic responders recruited and activated by microglia are neutrophils.18,19 Neutrophils comprise a large fraction of leukocytes in murine blood (7.4 – 33.9%, “young” C57BL/6J mice20). They assist in maintaining tissue homeostasis by neutralizing foreign agents, regulating the immune response, and phagocytosing dead tissue21,22. Under inflammatory conditions, the spatiotemporal interplay between microglia and neutrophils is poorly understood. A missed window of interaction is highly problematic in histological study where a single time point reveals a snapshot of the temporally complex immune response, which changes dynamically over time. Here, we use in vivo imaging to overcome these constraints.

Earlier version: Documenting immune cell interactions in the retina over time has been challenged by insufficient resolution and contrast to visualize single cells in the living eye. The microscopic size of immune cells requires exceptional resolution for detection. Recently, advances in AOSLO imaging have provided micron-level resolution and enhanced contrast for imaging individual immune cells in the retina and without requiring extrinsic dyes7,23. AOSLO provides multi-modal information from confocal reflectance, phase-contrast and fluorescence modalities, which can reveal a variety of cell types simultaneously in the living eye. Here, we used confocal AOSLO to track changes in reflectance at cellular scale. Phase-contrast AOSLO provides detail on highly translucent retinal structures such as vascular wall, single blood cells27–29, PR somata30, and is well-suited to image resident and systemic immune cells.7,23 Fluorescence AOSLO provides the ability to study fluorescently-labeled cells25,31,32 and exogenous dyes27,33 throughout the living retina. These modalities used in combination have recently provided detailed images of the retinal response to a model of human uveitis.23,34 Together, these innovations now provide a platform to visualize, for the first time, the dynamic interplay between many immune cell types, each with a unique role in tissue inflammation. We combine these innovative modalities with conventional histology as well as commercial OCT to reveal the progressive nature of the cellular response to acute retinal injury.

Revised version: Documenting immune cell interactions in the retina over time has been challenged by insufficient resolution and contrast to visualize single cells in the living eye. The microscopic size of immune cells requires exceptional resolution for detection. Recently, advances in AOSLO imaging have provided micron-level resolution and enhanced contrast for imaging individual immune cells in the retina without requiring extrinsic dyes7,23. AOSLO provides multi-modal information from confocal reflectance, phase-contrast and fluorescence modalities, which can reveal a variety of cell types simultaneously in the living eye. Here, we used confocal AOSLO to track changes in reflectance at cellular scale. Phase-contrast AOSLO provides detail on highly translucent retinal structures such as vascular wall, single blood cells27–29, PR somata30, and is well-suited to image resident and systemic immune cells.7,23 Fluorescence AOSLO provides the ability to study fluorescently-labeled cells25,31,32 and exogenous dyes27,33 throughout the living retina. These modalities used in combination have recently provided detailed images of the retinal response to a model of human uveitis.23,34 Together, these innovations now provide a platform to visualize, for the first time, the dynamic interplay between many immune cell types, each with a unique role in tissue inflammation. We combine these innovative modalities with conventional histology and commercial SLO/OCT to reveal the progressive nature of the cellular response to acute retinal injury.

Here, we ask the question: “To what extent do microglia/neutrophils respond to acute neural loss in the retina?” To begin unraveling the complexities in this response, we deploy a deep retinal laser ablation model. Using AOSLO, we track and characterize the changes in microglia, neutrophils and retinal structure within hours, days and months after acute laser exposure.

Methods

Removed in the revised version: Key Resources Table

Mice

All experiments herein were approved by the University Committee on Animal Resources (protocol #: UCAR-2010-052E) and according to the Association for Research in Vision and Ophthalmology statement for the Use of Animals in Ophthalmic and Vision Research as well as institutional approvals by the University of Rochester. C57BL/6J (#000664, Jackson Labs, Bar Harbor, ME) mice were used to track the retinal phenotype after laser exposure.

Heterozygous CX3CR1-GFP (#005582, Jackson Labs) mice were used to track GFP-expressing microglia. Mice with heterozygous transgenic expression of tdTomato in neutrophils (“Catchup” mice) were provided by the foundry lab of M. Gunzer.35 18 mice (10 male, 8 female) in total were used for in vivo imaging. 10 additional mice (5 male, 5 female) were used for ex vivo histology. Age for all mice used for this work were post-natal weeks 6 - 24.

Preparation for in vivo imaging

Earlier version: Mice were anesthetized with intra-perotineal injection of ketamine (100 mg/kg) and xylazine (10 mg/kg). The pupil was dilated with 1% tropicamide (Sandoz, Basel, Switzerland) and 2.5% phenylephrine (Akorn, Lake Forest, IL). A custom contact lens (1.5 mm base curve, 3.2 mm diameter, +10 diopter power, Advanced Vision Technologies, Lakewood, CO) was fitted to the eye. For a subset of experiments, 50 mg/kg fluorescein (AK-FLUOR, Akorn, Decatur, IL) was administered by intra-peritoneal injection to confirm vascular integrity and perfusion status after injury. During AOSLO imaging, anesthesia was supplemented with 1% (v/v) isoflurane in oxygen and mice were maintained at 37°C via electric heat pad. Eye hydration was maintained throughout imaging with regular application of saline eye drops (Refresh tears, Allergan, Sydney, Australia) and lubricating eye gel (Genteal, Alcon Laboratories Inc. Fort Worth, TX). Mice were placed in a positioning frame with 6-degrees of freedom to allow for stable animal positioning and aid in retinal navigation relative to the beam of the AOSLO and OCT.

Revised version: Mice were anesthetized with intra-perotineal injection of ketamine (100 mg/kg) and xylazine (10 mg/kg). The pupil was dilated with 1% tropicamide (Sandoz, Basel, Switzerland) and 2.5% phenylephrine (Akorn, Lake Forest, IL). A custom contact lens (1.5 mm base curve, 3.2 mm diameter, +10 diopter power, Advanced Vision Technologies, Lakewood, CO) was fitted to the eye. For a subset of experiments, 50 mg/kg fluorescein (AK-FLUOR, Akorn, Decatur, IL) was administered by intra-peritoneal injection to confirm vascular integrity and perfusion status after injury. During AOSLO imaging, anesthesia was supplemented with 1% (v/v) isoflurane in oxygen and mice were maintained at 37°C via electric heat pad. Eye hydration was maintained throughout imaging with regular application of saline eye drops (Refresh tears, Allergan, Sydney, Australia) and lubricating eye gel (Genteal, Alcon Laboratories Inc. Fort Worth, TX). Mice were placed in a positioning frame with 6-degrees of freedom to allow for stable animal positioning and aid in retinal navigation.

In vivo AOSLO imaging

Earlier version: Four light sources were used for AOSLO imaging. A 904 nm diode source (12 µW, Qphotonics, Ann Harbor, MI) was used for wavefront sensing. A second 796 nm superluminescent diode (196 µW, Superlum, Cork, Ireland) was used for reflectance imaging, including confocal and phase contrast modes.27,31 A third 488 nm light source (56 µW, Toptica Photonics, Farmington, NY) was used to visualize GFP-positive microglia in CX3CR1-GFP mice. A fourth 561 nm light source (95 µW, Toptica Photonics) was used to visualize tdTomato-positive neutrophils in Catchup mice. All light sources were fiber-coupled and axially combined through the AOSLO system31. Fast (15.4 kHz) and slow (25 Hz) scanners create a raster scan pattern, which is relayed through a series of afocal telescopes to and from the eye. A Shack-Hartmann wavefront sensor (consisting of a lenslet array and a Rolera XR camera, QImaging, Surrey, Canada) measures the aberrations of the eye and a deformable mirror (ALPAO, Montbonnot-Saint-Martin, France) provided the wavefront correction. Reflected 796 nm light was collected with a photomultiplier tube (H7422-50, Hamamatsu Photonics, Hamamatsu, Japan). All confocal reflectance images were captured with a 30 µm pinhole (1.3 Airy Disc Diameters, ADD). Phase-contrast was achieved by displacing the pinhole relative to the principal axis of the detection plane as previously described.25 Fluorescence was captured with a photomultiplier tube (H7422-40, Hamamatsu) either coupled with a 520Δ35 band-pass filter (FF01-520/35-25, Semrock, Rochester, NY) for GFP emission, or a 630Δ92 band-pass filter (FF01-630/92-25, Semrock) for tdTomato emission. All fluorescent images were captured with a confocal 50 µm pinhole (2.1 ADD).

Revised version: Four light sources were used for AOSLO imaging. A 904 nm diode source (12 µW, Qphotonics, Ann Harbor, MI) was used for wavefront sensing. A second 796 nm superluminescent diode (196 µW, Superlum, Cork, Ireland) was used for reflectance imaging, including confocal and phase contrast modes.27,31 A third 488 nm light source (56 µW, Toptica Photonics, Farmington, NY) was used to visualize GFP-positive microglia in CX3CR1-GFP mice. A fourth 561 nm light source (95 µW, Toptica Photonics) was used to visualize tdTomato-positive neutrophils in Catchup mice. All light sources were fiber-coupled and axially combined through the AOSLO system31. Fast (15.4 kHz) and slow (25 Hz) scanners create a raster scan pattern, which is relayed through a series of afocal telescopes to and from the eye. A Shack-Hartmann wavefront sensor (consisting of a lenslet array and a Rolera XR camera, QImaging, Surrey, Canada) measures the aberrations of the eye and a deformable mirror (ALPAO, Montbonnot-Saint-Martin, France) provided the wavefront correction. Reflected 796 nm light was collected with a photomultiplier tube (H7422-50, Hamamatsu Photonics, Hamamatsu, Japan). All confocal reflectance images were captured with a 30 µm pinhole (1.3 Airy Disc Diameters, ADD). Phase-contrast was achieved by displacing the pinhole relative to the principal axis of the detection plane as previously described.25 Fluorescence was captured with a photomultiplier tube (H7422-40, Hamamatsu) either coupled with a 520Δ35 band-pass filter (FF01-520/35-25, Semrock, Rochester, NY) for GFP emission, or a 630Δ92 band-pass filter (FF01-630/92-25, Semrock) for tdTomato emission. All fluorescent images were captured with a confocal 50 µm pinhole (2.1 ADD). Image field sizes were either 4.98° x 3.95° or 2.39° x 1.94°. NIR and visible imaging channels were made coplanar, compensating for longitudinal chromatic aberration by independently focusing each light source onto the same axial structure in the retina. Through-focus stacks were acquired by sequentially changing the focus from NFL to the PR outer segments by using the defocus term on the custom adaptive optics control software.

Removed in the revised version: Image field sizes were either 4.98° x 3.95° or 2.39° x 1.94°. NIR and visible imaging channels were made coplanar, compensating for longitudinal chromatic aberration by independently focusing each light source onto the same axial structure in the retina. Through-focus stacks were acquired by sequentially changing the focus from NFL to the PR outer segments by using the defocus term on the custom adaptive optics control software.

As described previously, red blood cell (RBC) imaging was achieved by combining phase-contrast imaging with a strategy to arrest the slow galvanometer scanner and let the resonant scanner project a single “line” (0.71° scan angle) on the retina. This enabled RBC flux imaging by positioning this line orthogonal to the direction of flow for single capillaries. As blood cells moved through capillaries, they were “self-scanned” producing images of RBCs in space/time27.

PR laser damage model and post-injury time points for imaging

488 nm light (Continuous wave Laser diode, ±4 nm bandwidth, Toptica Photonics) was used to create an acute laser injury. 785 µW of 488 nm light was projected through the AOSLO and focused onto the PR outer segments (Figure 1a) for 3 minutes in a single line on the retina subtending 24 x 1 µm to concentrate the power to a small region. Laser injuries were placed between 5-15° from the optic disc. To avoid absorption confounds, we refrained from placing lesions beneath large retinal vessels. For experiments examining neutrophil involvement, lesions were placed <100 µm away from retinal veins and within microns of capillaries to increase the chance of extravasation through these preferred pathways in the retina36–38. As many as four such lesions were placed per retina. This protocol produced a hyper-reflective phenotype in the >40 locations across 28 mice. In rare cases, the exposure yielded no hyper-reflective lesion and were often in mice with high retinal motion, where the light dosage was spread over a larger retinal area. These locations were not included in the in-vivo or histological analysis.

Laser injury assessed with commercial SLO and OCT.

Earlier version: (A) 488 nm light is focused onto the photoreceptor outer segments using AOSLO. Created with Biorender.com. (B) 30° SLO images of NIR reflectance, blue reflectance and fluorescein angiography of a mouse retina 1 day after laser exposure. Three focal planes are shown. NIR and blue reflectance reveal small hyperreflective regions below the superficial plane. Fluorescein reveals intact vasculature with no sign of leakage. Arrows indicate regions with imparted laser damage (1-4). (C) OCT B-scans passing through laser-exposed regions indicated in B. Exposures produced a focal hyperreflective band within the ONL with adjacent retina appearing healthy. OCT images were spatially averaged (∼30 µm, 3 B-scans). Scale bars = 200 µm horizontal, 200 µm vertical.

Revised version: (A) 488 nm light is focused onto the photoreceptor outer segments using AOSLO. Created with Biorender.com/gqjx0zy. (B) 30° SLO images of NIR reflectance, blue reflectance and fluorescein angiography of a mouse retina 1 day after laser exposure. Three focal planes are shown. NIR and blue reflectance reveal small hyperreflective regions below the superficial plane. Fluorescein reveals intact vasculature with no sign of leakage. Arrows indicate regions with imparted laser damage (1-4). (C) OCT B-scans passing through laser-exposed regions indicated in B. Exposures produced a focal hyperreflective band within the ONL with adjacent retina appearing healthy. OCT images were spatially averaged (∼30 µm, 3 B-scans). Scale bars = 200 µm horizontal, 200 µm vertical.

Added in the revised version: Lesion location tracked from minutes to 1 day with OCT.

Added in the revised version: After baseline OCT acquisition, OCT was performed every 5-7 minutes for one hour after 488 nm light exposure. 6 hour and 1 day time points were subsequently acquired. A band of hyperreflectivity forms near the OPL/ONL interface within 30 minutes of 488 nm light exposure. Hyperreflective band, spreads deeper into the ONL within ∼1 hour. OCT images were spatially averaged (∼30 µm, 8 B-scans). Scale bar = 40 µm horizontal, 100 µm vertical.

Earlier version: Throughout this work, we assessed the effects of the laser injury for the following time points: baseline/control, 1 day (defined as 18 – 28 hours post-damage), 3 days, 7 days and 2 months.

Revised version: Throughout this work, we assessed the effects of the laser injury for the following time points: baseline/control, 1 day (18 – 28 hours), 3 days, 7 days and 2 months.

In vivo SLO, and OCT imaging

Earlier version: To confirm global ocular health and changes imparted by the laser damage, a commercial Heidelberg Spectralis system (Heidelberg, Germany) was used to acquire scanning laser ophthalmoscopy (SLO) and optical coherence tomography (OCT) images. 30° and 55° fields were used for SLO acquisitions. The 30° field was used for OCT acquisitions. For some experiments, fluorescein angiography was captured by imaging the retina within 10 minutes of fluorescein administration (details above). The fluorescence mode of the SLO also enabled wide-field images of GFP positive microglia.

Revised version: To confirm global ocular health and changes imparted by the laser damage, a commercial Heidelberg Spectralis system (Heidelberg, Germany) was used to acquire SLO and OCT images. 30° and 55° fields were used for SLO acquisitions. The 30° field was used for OCT acquisitions. For some experiments, fluorescein angiography was captured by imaging the retina within 10 minutes of fluorescein administration (details above). The fluorescence mode of the SLO also enabled wide-field images of GFP positive microglia.

Earlier version: OCT was acquired with the following settings: ART 16 (automatic real-time tracking) was used to average B-scans and improve signal. To capture several damage locations in a single field, we used a coarse scan area (61 B-scans, 1.02 x 0.85 mm). A dense 3D data cube was also captured (49 B-scans, 513 x 171 µm) to provide detailed information regarding the axial nature of the laser damage. “Follow-up” mode, which allows the HRA software to return to the same retinal location, was used whenever possible. To reveal the cross-sectional profile for each lesion, several adjacent B-scans were spatially averaged (∼30 µm).

Revised version: OCT was used to provide detailed information regarding the axial nature of the laser damage. We used a coarse scan area to capture several damage locations in a single field (61 B-scans, 1.02 x 0.85 mm). A dense 3D data cube was also captured (49 B-scans, 513 x 171 µm). “Follow-up” mode, which allows the HRA software to return to the same retinal location, was used whenever possible. To reveal the cross-sectional profile for each lesion, several adjacent B-scans were spatially averaged (∼30 µm).

Preparation for ex vivo imaging

Earlier version: Mice were euthanized by CO2 asphyxiation followed by cervical dislocation. Within 5 minutes of asphyxiation, eyes were enucleated and placed in 4% paraformaldehyde (PFA, diluted from: #15714-S, Electron Microscopy Sciences, Hatfield, Pennsylvania) in 1x phosphate buffered saline (PBS, #806552, Sigma-Aldrich, St. Louis, MO) for one hour at room temperature. Eyes were dissected to remove the cornea, lens and vitreous. Each eye cup was placed in one well of a 24-well plate containing 0.5 ml of 0.8% PFA and left overnight at 4°C. The retina was separated from the retinal pigmented epithelium/choroid with attention to preserve inner retinal layers “up” orientation. If not applying antibody, the tissue was directly flat mounted (see below). For antibody staining, the retina was placed in 1x BD perm/wash buffer (#554723, BD Biosciences, Franklin Lakes, NJ) with 5% donkey serum (#D9663, Sigma-Aldrich) diluted in PBS for overnight incubation at room temperature with gentle shaking. The following was performed in the dark. Ly-6G-647 antibody (1:200, #127610, RRID:AB_1134159, BioLegend, San Diego, CA) and DAPI (1:500 of 10 mg/ml stock, #4083, Cell Signaling Technology, Danvers, MA) were diluted in 1x perm/wash buffer and retinas were incubated for 3 days at room temperature with gentle shaking.

Revised version: Mice were euthanized by CO2 asphyxiation followed by cervical dislocation. Within 5 minutes of asphyxiation, eyes were enucleated and placed in 4% paraformaldehyde (PFA, diluted from: #15714-S, Electron Microscopy Sciences, Hatfield, Pennsylvania) in 1x phosphate buffered saline (PBS, #806552, Sigma-Aldrich, St. Louis, MO) for one hour at room temperature. Eyes were dissected to remove the cornea, lens and vitreous. Each eye cup was placed in one well of a 24-well plate containing 0.5 ml of 0.8% PFA and left overnight at 4°C. The retina was separated from the retinal pigmented epithelium/choroid with attention to preserve inner retinal layers “up” orientation. If not applying antibody, the tissue was directly flat mounted (see below). For antibody staining, the retina was placed in 1x BD perm/wash buffer (#554723, BD Biosciences, Franklin Lakes, NJ) with 5% donkey serum (#D9663, Sigma-Aldrich) diluted in PBS for overnight incubation at room temperature with gentle shaking. The following was performed in the dark. Ly-6G-647 antibody (1:200, #127610, RRID:AB_1134159, BioLegend, San Diego, CA) and DAPI (1:500 of 10 mg/ml stock, #4083, Cell Signaling Technology, Danvers, MA) were diluted in 1x perm/wash buffer and retinas were incubated for 3 days at room temperature with gentle shaking. Retinas were washed with PBS three times over three hours. The retina was cut into 4 radially symmetrical petals, flat-mounted on a glass slide in Vectashield mounting buffer (H-1000-10, Vector labs, Newark, CA) with a #1.5 cover slip (#260406, Ted Pella Inc., Redding, CA), sealed with nail polish and stored at 4°C until imaged.

Added in the revised version: Ex vivo confocal imaging

Added in the revised version: Whole-mount retinas were imaged with a Nikon A1 confocal microscope (Melville, NY). DAPI (405 nm ex, 441Δ66 nm em), CX3CR1-GFP (488 nm ex, 525Δ50 nm em), and Catchup/anti-Ly-6G-647 (635 nm ex, 665Δ50 nm em) were simultaneously imaged. Z-stacks (0.1 or 0.5 µm step size) at control or laser-damaged locations were acquired with a 60x oil objective, producing images that were 295 x 295 µm. Larger (796 x 796 µm) z-stacks were acquired by blending several 60x z-stack acquisitions (3x3, 15% overlap) using Nikon NIS Elements software. z-stacks were re-sliced (Imagej FIJI42, X-Z dimension) to visualize fluorescence depth profiles.

Removed in the revised version: Retinas were washed with PBS three times over three hours. The retina was cut into 4 radially symmetrical petals, flat-mounted on a glass slide in Vectashield mounting buffer (H-1000-10, Vector labs, Newark, CA) with a #1.5 cover slip (#260406, Ted Pella Inc., Redding, CA), sealed with nail polish and stored at 4°C until imaged.

Endotoxin-Induced Uveitis protocol

To serve as a positive control and show evidence of known neutrophil invasion, we adopted the endotoxin-induced uveitis (EIU) model to confirm that fluorescent Catchup neutrophils could be observed in vivo with AOSLO. This was performed in two mice. The EIU model has been described previously.34,39 Briefly, we performed intravitreal injections of lipopolysaccharide (LPS, #L4391, Sigma-Aldrich). Mice were anesthetized with ketamine and xylazine. A 34 gauge Hamilton needle was used to deliver 1 µL (1 ng) of LPS diluted in PBS into the vitreous posterior to the limbus. 1 day post-LPS injection, mice were either imaged with AOSLO (Catchup mice) or collected for ex vivo histology (Ly-6G-647 stained C57BL/6J mice) to confirm fluorescent neutrophil presence in the neural parenchyma.

Removed in the revised version: Ex vivo confocal imaging

Removed in the revised version: Whole-mount retinas were imaged with a Nikon A1 confocal microscope (Melville, NY). DAPI (405 nm ex, 441Δ66 nm em), CX3CR1-GFP (488 nm ex, 525Δ50 nm em), and Catchup/anti-Ly-6G-647 (635 nm ex, 665Δ50 nm em) were simultaneously imaged. Z-stacks (0.1 or 0.5 µm step size) at control or laser-damaged locations were acquired with a 60x oil objective, producing images that were 295 x 295 µm. Larger (796 x 796 µm) z-stacks were acquired by blending several 60x z-stack acquisitions (3×3, 15% overlap) using Nikon NIS Elements software. z-stacks were re-sliced (Imagej FIJI42, X-Z dimension) to visualize fluorescence depth profiles.

AOSLO image processing

To correct residual motion from heart rate and respiration, AOSLO videos were registered with a custom cross-correlation-based frame registration software.40,41 Motion correction was also applied to simultaneously collected fluorescence videos. After registration, confocal, phase-contrast and fluorescence AOSLO videos were temporally averaged (250 frames, 10 seconds). Blood perfusion maps were computed by calculating the standard deviation of pixel intensity over 30 seconds32.

INL + ONL nuclei quantification

Raw DAPI z-stacks were used for manual counting of nuclei (n = 10 mice, 3 regions per time point). Analysis regions were circular, with a 50 µm diameter (corresponding to the size and shape of the damage region seen with confocal AOSLO) and analyzed in depth producing volumetric cylinders through the INL or outer nuclear layer (ONL, Figure 3d,e). Nuclei were counted manually using the “Cell Counter” plugin in Imagej (Author: Kurt De Vos, Imagej version 1.53q).

Added in the revised version: Laser damage temporally tracked with AOSLO and OCT.

Added in the revised version: Laser-exposed retina was tracked with OCT (A), confocal (B) and phase-contrast (C) AOSLO for baseline, 1, 3, 7 day and 2 month time points. OCT and confocal AOSLO display a hyperreflective phenotype that was largest/brightest at 1 day and became nearly invisible by 2 months. Dashed oval indicates region targeted for laser injury. Phase-contrast AOSLO revealed disrupted photoreceptor soma’s 1 day after laser injury. Phase contrast data was not acquired for remaining time points due to development of cataract which obscured the phase contrast signal. OCT images were spatially averaged (∼30 µm, 8 B-scans). Scale bars = 40 µm horizontal, 100 µm vertical.

Added in the revised version: Retinal histology confirms photoreceptor ablation and preservation of inner retinal cells.

Added in the revised version: Cross sectional view (A) and en-face (B+C) images of DAPI-stained whole-mount retinas at laser injury locations over time. By 1 day, ONL becomes thicker at lesion location, but thinner by 3 and 7 days. By 2 months, the ONL appeared similar to that of control. The inner (B, solid rectangle) and outer (C, dashed rectangle) stratum of ONL show axial differences in ONL loss. Most cell loss was seen in the outer aspect of the ONL (C). Scale bars = 40 µm. (D) Cross section of DAPI-stained retina displaying INL and ONL regions for quantification. Each analysis region was 50 µm across and encompassed the entire depth of the INL or ONL. (E) En-face images show 50 µm diameter circles used for analysis. (F) Nuclei density for control and post-injury time points. ONL nuclei were reduced at 3 and 7 days (p = 0.17 and 0.07 respectively) while INL density remained stable (n = 10 mice, 3 unique regions per time point). Error bars display mean ± 1 SD.

PR + Microglia nuclei volume quantification

Earlier version: A single DAPI-stained z-stack (OPL + ONL) from a CX3CR1-GFP mouse was used to quantify nuclear volume of PR’s (n = 20) and microglia (n = 14). Data was centered at a lesion location 3 days post-laser exposure. Imagej was used to manually measure the en-face diameter of nuclei in their short and long axis. We averaged the two measurements and assumed spherical shape for analysis. These measurements aided in the differentiation of PR’s, invading microglia and PR phagosomes.

Revised version: A single DAPI-stained z-stack (OPL + ONL) from a CX3CR1-GFP mouse was used to quantify nuclear volume of PR’s (n = 20) and microglia (n = 14). Data was centered at a lesion location 3 days post-laser exposure. Imagej was used to manually measure the en-face diameter of nuclei in their short and long axis. We averaged the two measurements and assumed spherical shape for analysis.

Added in the revised version: These measurements aided in the differentiation of PR’s, invading microglia and PR phagosomes.

Quantification of RBC flux in capillaries

Using the capillary line-scan approach27,43,44, we quantified RBC flux within an epoch of 1 second. This spanned several cardiac cycles.29,45 To improve SNR, space-time images were convolved with a Gaussian spatio-temporal filter (σ = 7.5 pixels, 0.33 µm). This strategy did not interfere with spatial resolution as pixels oversample the optical point-spread of the AOSLO by >20x.27 RBC flux was determined by manually marking blood cells using the “Cell counter” plugin in Imagej.

Neutrophil density quantification

Ly-6G-647-stained retinal tissue was used to manually count neutrophils within montaged 796 x 796 µm z-stacks. Cells were counted at control (n = 4 locations, 2 mice) and lesioned (n = 4 locations, 2 mice) locations using the “Cell Counter” Imagej plugin. All values are reported as mean ± SD. There were no data points omitted from any of the analysis reported in this work.

Results

Characterization of deep focal laser damage

Four complementary imaging modalities were used to evaluate the nature and localization of the focal laser damage induced by 488 nm light: wide field SLO, OCT, AOSLO and post-mortem histology. Each is reported in turn below.

Laser damage induces focal hyperreflective lesions in outer retina imaged with wide-field SLO

Earlier version: To observe global and focal retinal health, we used commercial SLO. 1 day post-488 nm light exposure, both NIR and blue reflectance modalities showed hyperreflective lesions, most apparent at a deeper retinal focus position. When SLO focus was brought to the inner retina, lesions were not visible and retina appeared healthy (Figure 1b). Despite deep retinal damage, fluorescein angiography did not reveal dye leakage (Figure 1b, bottom), indicating the blood retinal barrier remained intact.

Revised version: To observe global and focal retinal health, we used commercial SLO. 1 day post-488 nm light exposure, both NIR and blue reflectance modalities showed hyperreflective lesions, most apparent at a deeper retinal focus position. In the inner retina, lesions were not visible and the retina appeared healthy (Figure 1b). Despite deep retinal damage, fluorescein angiography did not reveal dye leakage (Figure 1b, bottom), indicating the blood retinal barrier remained intact.

OCT B-scans reveal outer retinal hyperreflection without inner retinal damage

Earlier version: To assist in determination of which retinal layers are damaged by the 488 nm laser, we used OCT. Within 30 minutes post-laser exposure, OCT revealed a zone of hyperreflection within the ONL (Figure 1 – figure supplement 1). 1 day post-lesion, the focal hyperreflection remained localized to the ONL (∼50 µm wide, Figure 1c). Retina outside of damage foci appeared normal. The hyperreflective phenotype persisted through 7 days and was cleared by 2 months (Figure 2a). There did not appear to be any cellular excavation, “cratering” or evidence of edema for any time point assessed. Despite the hyperreflective ONL phenotype, Bruch’s membrane appeared intact for all time points, evidenced by lack of fluorescein leakage at the site of lesion. Additionally, retinal vessels appeared normal in OCT B-scans.

Revised version: To assist in determination of which retinal layers are damaged by the 488 nm laser, we used OCT. Within 30 minutes post-laser exposure, OCT revealed a zone of hyperreflection within the ONL (Figure 1 – figure supplement 1). 1 day post-lesion, the focal hyperreflection remained localized to the ONL (∼50 µm wide, Figure 1c). Retina outside of damage foci appeared normal. The hyperreflective phenotype persisted through 7 days and was cleared by 2 months (Figure 2a). There did not appear to be any cellular excavation, “cratering” or evidence of edema for any time point assessed. Bruch’s membrane appeared intact for all time points, evidenced by lack of fluorescein leakage at the site of lesion. Additionally, retinal vessels appeared normal in OCT B-scans.

Removed in the revised version: Laser damage temporally tracked with AOSLO and OCT. Laser-exposed retina was tracked with OCT.

Removed in the revised version: (A), confocal (B) and phase-contrast (C) AOSLO for baseline, 1, 3, 7 day and 2 month time points. OCT and confocal AOSLO display a hyperreflective phenotype that was largest/brightest at 1 day and became nearly invisible by 2 months. Dashed oval indicates region targeted for laser injury. Phase-contrast AOSLO revealed disrupted photoreceptor soma’s 1 day after laser injury. Phase contrast data was not acquired for time points 3 days-2 months due to development of cataract which obscured the phase contrast signal. OCT images were spatially averaged (∼30 µm, 8 B-scans). Scale bars = 40 µm, OCT image vertical scale bar = 100 µm.

AOSLO reveals outer retinal damage with confocal and phase-contrast modalities

Confocal AOSLO provided micron-level detail of the lesioned area. Confirming OCT findings, we observed hyperreflective changes localized within the ONL that were, by 1 day, the brightest at the OPL/ONL interface (Figure 2b). At this plane, the lesions manifest as an ellipse with the long axis in the direction of the line scan used to create the lesion. The most prominent hyperreflective phenotype was seen at 1 day post lesion, with diminishing size and brightness by day 3 and 7. The phenotype was largely diminished by 2 months (Figure 2b).

Using phase-contrast AOSLO also allowed visualization of translucent cells within the retina which enabled us to image PR somas of the ONL25,30. Normally, the ONL is comprised of PR somata that are densely-packed30,46. We found that the dense packing of individual somata was disrupted 1 day post-exposure (∼50 µm ovoid, Figure 2c), suggesting degradation or ablation of the cell membrane.

Confirmation of outer retinal cell loss using post-mortem histology

To assess the extent of cell loss caused by focal laser exposure, we performed DAPI staining on whole-mount retinal tissue using the same time points assessed for in vivo imaging. 1 day after laser exposure, we observed mild thickening of the ONL compared to unexposed locations only microns away. At 3 and 7 days, local ONL thinning was observed (Figure 3a). En-face planes within the ONL revealed a loss of PR nuclei in the outer aspect of the ONL for 3 and 7 day time points, while inner ONL exhibited little evidence of cell loss (Figure 3b, c), illustrating the precise axial confinement of the laser damage induced by this method. By 2 months, the lesion’s overall appearance and ONL thickness returned to baseline (Figure 3a-c). This histological finding corroborates the OCT findings observed in vivo.

Removed in the revised version: Retinal histology confirms photoreceptor ablation and preservation of inner retinal cells.

Removed in the revised version: Cross sectional view (A) and en-face (B+C) images of DAPI-stained whole-mount retinas (5 mice) at laser injury locations over time. By 1 day, ONL becomes thicker at lesion location, but thinner by 3 and 7 days. By 2 months, the ONL appeared similar to that of control. The inner (B, solid rectangle) and outer (C, dashed rectangle) stratum of ONL show axial differences in ONL loss. Most cell loss was seen in the outer aspect of the ONL (C). Scale bars = 40 µm. (D) Cross section of DAPI-stained retina displaying INL and ONL regions for quantification. Each analysis region was 50 µm across and encompassed the entire INL or ONL. (E) En-face images show 50 µm diameter circles used for analysis. (F) Nuclei density for control and post-injury time points. ONL nuclei were reduced at 3 and 7 days (p = 0.17 and 0.07 respectively) while INL density remained stable (n = 10 mice, 3 unique regions per time point). Error bars display mean ± 1 SD.

ONL nuclear counts were reduced at locations of laser exposure. In comparison to control locations (501317 nuclei/mm2 ± 46198 nuclei/mm2, mean ±SD, similar to previous work47), at 3 and 7 days post-exposure, we found a corresponding reduction in the density of ONL nuclei of 18% (408965 nuclei/mm2 ± 25621 nuclei/mm2) and 27% (367542 nuclei/mm2 ± 9038 nuclei/mm2) respectively. Compared to control locations, the 3 and 7 day data resulted in p-values of 0.17 and 0.07, respectively (Student’s paired two-tailed t-test). 2 months after damage, PR nuclear densities were indistinguishable from that of control (Figure 3f, gray bars). Despite losses of PR nuclei, total cell nuclei within the INL remained unchanged for all time points (Figure 3f, black bars). These data indicate that the laser exposure focally ablated PR’s whilst leaving inner retinal cells intact.

Retinal vasculature unaffected by deep retinal lesion

A concern with laser lesions of this type is that it may coagulate retinal vessels. Motion contrast26,27 images revealed that the vasculature remained perfused from hours to months after laser damage suggesting that acute or long term changes are not imparted by the laser injury (Figure 4). None of the primary vascular stratifications within the NFL, IPL and OPL of the mouse retina32 showed stopped flow as a result of the laser lesion, reinforcing the findings above regarding the axial confinement of the damage to the outer retina.

Added in the revised version: Motion-contrast images reveal vascular perfusion status in response to laser damage.

Removed in the revised version: Microvascular perfusion unchanged after laser damage.

Earlier version: A single location tracked over time and at three vascular plexuses using AOSLO. Motion-contrast images were generated from confocal videos to reveal vascular perfusion status. Retinal vasculature remained perfused for all time points tracked and at all depths. White oval indicates damage location. Scale bar = 40 µm.

Revised version: A single location was tracked over time at three vascular plexuses using AOSLO. Retinal vasculature remained perfused for all time points tracked and at all depths. White oval indicates damage location. Scale bar = 40 µm.

Added in the revised version: Measurement of single-cell blood flux after laser damage using phase contrast AOSLO.

Added in the revised version: Mouse 1: (A) The vascular plexus corresponding to IPL (cyan) and OPL (red) were targeted for flux determination. Blood cell flux was measured for 2 capillaries within the same field, at different depths. Arrows show the location for repeated line scan acquisitions. Created with Biorender.com/hsfjbxh. (B) RBC flux images acquired up to 7 days post-damage. Scale bars = 10 ms horizontal, 5 µm vertical. (C) Capillary flux quantified over 7 days. Despite the outer capillary displaying higher flux, both inner and outer capillaries changed synchronously for each time point. (D) Correlation of inner and outer capillary flux. Linear regression model displays a weak positive correlation (black dotted line). Mouse 2: (E) Left: Representative 55° SLO image showing regions targeted for capillary flux measurement. One region was subject to 488 nm laser damage and the other was left unlasered (Control). Scale bar = 200 µm. Right: Capillaries targeted for blood cell flux measurement. Arrows show the location for repeated line scan acquisitions. Scale bar = 40 µm. (F) RBC flux images acquired up to 2 months post-damage. Scale bars = 10 ms horizontal, 5 µm vertical. (G) Capillary flux quantified over 2 months. Flux remained similar at lesion and control locations for all time points assessed. Gray shaded regions indicate the range for normal capillary flux in the healthy C57BL/6J mouse.48 (H) Correlation of flux in lesion and control locations. Linear regression model displays a positive correlation (black dotted line).

In addition to perfusion status, phase-contrast AOSLO also permitted analysis of single blood cell flux within capillaries27,48. We tracked capillary flux at different retinal depths within and above the lesion (IPL, OPL, 1 capillary each, Figure 4 – figure supplement 1a). Blood cell flux for capillaries within lesion locations were within the range of normal flux for the C57BL/6J mouse48 (Figure 4 – figure supplement 1b, c). Flux tracked from hours to days in these capillaries changed synchronously, displaying positive linear correlation (R2 = 0.59, Figure 4 – figure supplement 1c, d). This suggested any such changes in flux were a property of systemic perfusion, rather than locally imparted changes in flow due to the lesion. As an additional control, we evaluated blood flux in two distant capillaries (control, Figure 4 – figure supplement 1e). Both capillaries displayed similar flux values from minutes to 2 months post-injury (Figure 4 – figure supplement 1f, g) resulting in positive linear correlation (R2 = 0.78, Figure 4 – figure supplement 1h). Taken together, these findings suggest lesions do not appreciably impact local RBC delivery in the capillary network.

PR laser injury promotes a robust response in nearby microglia

To observe the microglial response to PR laser injury, we imaged fluorescent microglia in CX3CR1-GFP mice with both SLO and AOSLO. 1 day after injury, SLO revealed bright, focal congregations of microglia at injury locations in contrast to undamaged locations, which maintained a distribution of lateral tiling (Figure 5a, b). The global visualization of microglia was augmented by high resolution fluorescence AOSLO, providing enhanced detail of the microglial response to laser injury7,49. Whereas AOSLO imaging of microglia in the healthy retina displayed a distributed array of microglia with ramified processes (Figure 5c, left), laser damaged locations showed a congregation of cells 1 day post-injury with less lateral ramification (Figure 5c, right). Within hours of the laser exposure, we did not observe a photo-bleaching or death of regional microglia suggesting that while the laser exposure was sufficient to damage PR’s, it left retinal microglia intact.

Microglial response 1 day after laser injury imaged in vivo with fluorescence SLO and AOSLO.

(A) Left: Deep-focus NIR SLO fundus image (55° FOV) of laser-injured retina. White arrowheads point to damaged locations showing hyperreflective regions. Inset scale bar = 40 µm. Right: Fluorescence fundus image from same location. Fluorescent CX3CR1-GFP microglia are distributed across the retina and show congregations at laser-damaged locations. Scale bar = 200 µm. (B) Magnified SLO images of microglia at laser-damaged and control locations (indicated in A, right, white boxes). Control location displays distributed microglial whereas microglia at the lesion location are bright and focally aggregated. (C) Fluorescence AOSLO images show greater detail of cell morphology at the same scale. In control locations microglia showed ramified morphology and distributed concentration whereas damage locations revealed dense aggregation of many microglia that display less ramification. Scale bars = 40 µm.

With phase-contrast imaging targeting the ONL, we documented a rare event of putative pseudopod extension at a lesion site (Figure 5 – video 1). Given the axial complexity of microglia in this layer, it is now possible that microglial process dynamics may be revealed with this label-free approach.

We returned to the same laser-damaged locations to capture microglial appearance at baseline, 1, 3, 7 day and 2 month time points with AOSLO to track the natural history of the microglial response to PR damage. At 1 day post-injury, damage locations displayed aggregations of microglia. The surrounding microglia displayed process polarization with extensions projecting toward the injury (Figure 6). By days 3 and 7, microglia exhibited fewer lateral projections and somas have migrated into the ONL where they do not normally reside (Figure 6). By 2 months post-injury, the hyperreflective phenotype was absent and microglia once again occupied only the inner retina. The microglial distribution at 2 months was similar to baseline, with cells exhibiting radially symmetrical branching projections, similar to those prior to laser injury (Figure 6).

Microglial response to laser injury tracked with AOSLO.

Simultaneously acquired NIR confocal and fluorescence AOSLO images across different retinal depths. Data are from one CX3CR1-GFP mouse tracked for 2 months. Microglia swarm to hyperreflective locations within 1 day. Microglia maintain an aggregated density for days and resolve by 2 months after damage. Scale bar = 40 µm.

Added in the revised version: Hyperreflective appearance emerges before microglia swarm to damage location.

Added in the revised version: AOSLO confocal and fluorescence images were acquired for baseline, 30, 90 minute and 1 day post-laser exposure. The hyperreflective phenotype appeared within 30 minutes but microglia were not found to aggregate until 1 day post damage. Scalebar = 40 µm.

A standing question in OCT/confocal AOSLO lesion interpretation is whether microglia contribute or directly produce the hyperreflective phenotype seen in axial B-scans and en-face fundus images50,51. With confocal AOSLO, we find that the hyperreflective phenotype is visible as early as 30 minutes, becoming larger and brighter by 90 minutes. During these time points, simultaneously imaged microglia remained ramified and maintained a tiled arrangement, indicating that microglial are not the initial source of the lesion-induced hyperreflective appearance (Figure 6 – figure supplement 1).

Despite robust microglial involvement, neutrophils do not extravasate

While we observed a robust microglial response to PR damage, there was no evidence of neutrophil involvement at the times we examined (1, 3, 7 day, and 2 month follow-up).

In vivo fluorescence imaging allowed us to track neutrophils with AOSLO in Catchup mice35. In healthy mice, we observed a sparse population of circulating neutrophils flowing quickly within the largest retinal vessels (Figure 7 – video 1). We also observed neutrophils moving through single capillary branches (Figure 7 – video 2). We found that neutrophils show deformation within the small confines of the capillary lumen, often resulting in tube or pill-shaped morphology seen both in vivo and ex vivo (Figure 7b, top).

Neutrophil morphology imaged in vivo using AOSLO.

Earlier version: (A) Phase-contrast, motion-contrast and fluorescence AOSLO reveal the impact of passing neutrophils on single capillaries. A rare and exemplary event shows a neutrophil transiently impeding capillary blood flow for minutes in healthy retina. Scale bar = 40 µm. (B) In-vivo AOSLO and ex-vivo fluorescence microscopy show neutrophils in two states. Neutrophils within capillaries displayed elongated, tubular morphology. Extravasated neutrophils were more spherical. Bottom images show extravasated neutrophils in response to LPS model for comparison (not laser damage model). Scale bar = 20 µm.

Revised version: (A) Phase-contrast, motion-contrast and fluorescence AOSLO reveal the impact of passing neutrophils on single capillaries. A rare and exemplary event shows a neutrophil transiently impeding capillary blood flow for minutes in healthy retina. Scale bar = 40 µm. (B) In-vivo AOSLO and ex-vivo fluorescence microscopy show neutrophils in two states. Neutrophils within capillaries displayed elongated, tubular morphology. Extravasated neutrophils were more spherical. Bottom images show extravasated neutrophils in response to an EIU model for comparison (not laser damage model). Scale bar = 20 µm.

Added in the revised version: Positive control EIU model: wide-field image of ex-vivo neutrophils 1 day post-LPS injection.

Added in the revised version: Retinal whole mount (C57BL/6J mouse) stained for DAPI (left), Ly-6G-647 (middle) and merged (right) shows a large neutrophil response, many of which have extravasated into the retinal parenchyma. White box indicates region cropped and displayed in figure 7b. Scale bar = 40 µm.

Earlier version: After laser lesion, we found no evidence of neutrophil aggregation or extravasation for any time point assessed (Figure 8). There was also a notable lack of rolling/crawling neutrophils (or any putative leukocyte) in large arterioles or venules surrounding the injury (Figure 8 – video 1). Neutrophils closest to the ONL lesion were occasionally detected within the deepest retinal capillaries.

Revised version: After laser lesion, we found no evidence of neutrophil aggregation or extravasation for any time point assessed (Figure 8, Figure 8 – figure supplement 1). There was also a notable lack of rolling/crawling neutrophils (or any putative leukocyte) in large arterioles or venules surrounding the injury (Figure 8 – video 1). Neutrophils closest to the ONL lesion were occasionally detected within the deepest retinal capillaries. However, these neutrophils stayed within the retinal vasculature, as evidenced by their pill-shaped morphology and passage routes that follow the known vascular paths seen in confocal/phase contrast images (Figure 8 – video 2). Leukocytes, including neutrophils, often impede flow due to their large size (13.7 µm52) as they compress through capillaries <7 µm in diameter48. However, we found no evidence of permanently stopped flow, as rare stalls would re-perfuse similar to those previously characterized in healthy mice (Figure 7a, Figure 8 – video 3).48

Added in the revised version: Neutrophil response to laser injury tracked with AOSLO.

Removed in the revised version: Retinal damage location tracked with AOSLO does not show accumulation of neutrophils.

Earlier version: A single retinal location was tracked in a Catchup mouse from baseline to 2 months after lesion. Location of the lesion is apparent at 1 and 3 days post injury with diminishing visibility after 1 week. Despite the ability to detect neutrophils (Figure 7), we did not observe stalled, aggregated or an accumulation of neutrophils at any time point evaluated. This evaluation was confirmed at multiple depths ranging from the NFL to the ONL of the retina in vivo. Scale bar = 40 µm.

Revised version: A single retinal location was tracked in a Catchup mouse from baseline to 2 months after lesion. Location of the lesion is apparent at 1 and 3 days post injury with diminishing visibility after 1 week. We did not observe stalled, aggregated or an accumulation of neutrophils at any time point. This evaluation was confirmed at multiple depths ranging from the NFL to the ONL. Scale bar = 40 µm.

Added in the revised version: Acute neutrophil response to laser injury tracked with AOSLO.

Added in the revised version: A single mouse was tracked for 0.5 - 3.5 hours after a deep retinal lesion was placed at the center of the imaged field, adjacent to a large venule. Neutrophils did not extravasate within this early post-lesion window (*also see figure 8 – video 3).

Removed in the revised version: However, these neutrophils stayed within the retinal vasculature, as evidenced by their pill-shaped morphology and passage routes that follow the known vascular paths seen in confocal/phase contrast images (Figure 8 – video 2).

Removed in the revised version: Leukocytes, including neutrophils, often impede flow due to their large size (13.7 µm52) as they compress through capillaries <7 µm in diameter48. However, we found no evidence of permanently stopped flow, as rare stalls would re-perfuse similar to those previously characterized in healthy mice (Figure 7a).48

Earlier version: As a positive control, we used the EIU model in Catchup mice to show that we could image extravasated neutrophils. This model is known to induce a strong neutrophil response. With fluorescence AOSLO, we found Catchup-positive neutrophils within the retinal parenchyma, providing confirmation that extravasated neutrophils can be imaged with this modality. Moreover, we found that neutrophils that have extravasated into the retinal parenchyma tended to have a more spherical morphology rather than the compressed, pill-shaped morphology of neutrophils within capillaries (Figure 7b).

Revised version: As a positive control, we used the EIU model to show that we could image extravasated neutrophils. This model is known to induce a strong neutrophil response. With fluorescence AOSLO (Figure 7 – video 2) and ex-vivo confocal microscopy (Figure 7 – supplemental figure 1), we found an abundance of neutrophils within the retinal parenchyma 1-day after LPS injection. This confirms the validity of our experimental paradigm and that extravasated neutrophils can be imaged with these modalities. Moreover, we found that neutrophils that have extravasated into the retinal parenchyma tended to have a more spherical morphology rather than the compressed, pill-shaped morphology of neutrophils within capillaries (Figure 7b).

Ex vivo analysis confirms in vivo findings

To confirm our in vivo findings, we examined fluorescent microglia and neutrophils in laser-damaged retinal whole-mounts imaged with confocal microscopy. The progressive nature of the microglial response to PR damage and general lack of neutrophil participation corroborated in vivo findings.

Ex vivo: Microglia display dynamic morphological changes in lesion areas

Earlier version: Without laser damage, microglia exhibited a tiled distribution and stellate morphology with highly-ramified branching patterns (Figure 9a). 1 day after laser exposure, microglial somas aggregated to the lesion location. They began to migrate into the ONL and changed from the ramified, and pancake-flat morphology seen in the healthy retina, to a dagger-like axial morphology (figure 9b). At 3 and 7 days after lesion, microglia migrated deeper into the outer retina and remained aggregated with an axially elongated phenotype. 2 months after lesion, microglial somas were no longer found in the ONL. Instead, microglia redistributed similar to that of the healthy retina, once again co-stratifying predominantly with the NFL and plexiform layers of the retina (Figure 9a, b, Video 1).

Revised version: Without laser damage, microglia exhibited a tiled distribution and stellate morphology with highly-ramified branching patterns (Figure 9a). 1 day after laser exposure, microglial somas aggregated to the lesion location. They began to migrate into the ONL and changed from the ramified morphology seen in the healthy retina, to a dagger-like axial morphology (figure 9b). At 3 and 7 days after lesion, microglia migrated deeper into the outer retina and remained aggregated with an axially elongated phenotype. 2 months after lesion, microglial somas were no longer found in the ONL. Instead, microglia redistributed similar to that of the healthy retina, once again co-stratifying predominantly with the NFL and plexiform layers of the retina (Figure 9a, b, Video 1).

Earlier version: Neutrophil behavior under laser injury, as observed through ex vivo confocal microscopy, remains unresponsive despite marked microglial activation.

Revised version: Neutrophil and microglial behavior after laser injury, as observed through ex vivo confocal microscopy.

(A) En-face max intensity projection images of inner and outer (separated by approximate INL center) retinal microglia/neutrophils in Ly-6G-647-stained CX3CR1-GFP retinas. Microglia display focal aggregation in the outer retina for 1, 3 and 7 day time points that is resolved by 2 months. Neutrophils do not aggregate or colocalize to the injury location at any time point. Z-stacks were collected from 5 mice for the indicated time points. (B) Cross-sectional views of en-face z-stacks presented in A, including DAPI nuclear label. White dotted line indicates 100 µm region expanded below. Microglia migrate into the ONL by 1, 3 and 7 days post-laser injury and return to an axial distribution similar to that of control by 2 months. The few neutrophils detected remained within the inner retina. Scale bars = 40 µm. (C) Orthogonal view of DAPI-stained retina with Ly-6G-647-labelled overlay 1 day post-laser-injury. In a rare example, 2 neutrophils are found within the IPL/OPL layers despite a nearby outer retinal laser lesion. Scale bar = 20 µm. (D) Magnified 3D cubes representing cell 1 and 2 in C. Cell 1 displays pill-shaped morphology and cell 2 is localized to a putative capillary branch-point. Each are confined within vessels suggesting they do not extravasate in response to laser injury.

Added in the revised version: Neutrophil/microglial response to laser injury tracked with ex vivo confocal microscopy.

Added in the revised version: Simultaneously acquired GFP-positive microglia and Ly-6G-647-positive neutrophils were imaged with confocal microscopy in 5 CX3CR1-GFP mice. En-face images for several retinal depths are displayed. By 1, 3 and 7 days post-lesion, microglia have migrated into the outer retina, many appearing amoeboid and displaying fewer laterally-branching projections. Despite the deep microglial response, neutrophils stay within the inner retina and are not found in the avascular outer retinal layers. Scale bar = 40 µm.

Added in the revised version: Microglial PR phagosomes in the outer retina assessed with ex vivo confocal imaging.

Added in the revised version: (A) En-face images of outer ONL in a DAPI-stained CX3CR1-GFP mouse 3 days post-laser-injury (top row). Microglia have infiltrated deep into the ONL and several PR phagosomes were identified. White arrows indicate locations for a single microglia (i), PR (ii) and PR phagosome (iii). These locations were expanded and displayed below. Microglia exhibited a heterogeneous nuclear staining pattern while PR nuclei exhibited homogenous DAPI fluorescence pattern. PR’s displayed this pattern regardless of whether they were within a microglial phagosome or not. Top scale bar = 20 µm, bottom scale bar = 2 µm. (B) A finely-sliced (0.1 µm step size) outer retinal z-stack of DAPI-stained CX3CR1-GFP retina was used to quantify the average nuclear volume for infiltrated microglia (n = 14 nuclei) and PR’s (n = 20 nuclei) for the same lesion site presented in A. On average, microglia had a statistically significant (p < 0.001) nuclear volume that was >3x that of PR’s. These measurements allowed us to discriminate microglial somas from PR phagosomes. Error bars display mean + 1 SD. (C) Cross sections of DAPI-stained outer retina in CX3CR1-GFP mice for 1, 3 and 7 days post-laser injury (n=3 mice). 3 representative planes (X-Z) through the lesion are displayed for each time point. Microglia form PR phagosomes within the ONL and microglial processes were seen extended into the PR inner/outer segment layer. Arrows label various morphological features seen at lesion sites: microglial somas (yellow), diving microglial process (violet), PR phagosome (red), microglial inner/outer segment process (cyan). Scale bar = 20 µm.

Whereas we found a robust response of deep microglia at the OPL to PR injury, we observed little migratory response of the microglia in the NFL or IPL (Figure 9 + Figure 9 – figure supplement 1) suggesting axial and lateral constraints on the extent of microglial recruitment.

Microglia form PR-containing phagosomes

A detailed examination of microglial involvement with PR somata in response to PR injury revealed unique microglia-PR interactions within the ONL. In cross-section, PR cells comprise a large portion of the retina reaching from the outer-segment tip, inner segments, somata and spherule/pedicle synaptic contacts at the OPL. Likewise, we saw microglial involvement with all of these layers. Within 1 day, we found microglial processes interspersed within the dense aggregation of PR somata within the ONL. Amoeboid cells enveloped the somata of PR’s and phagocytosis was detected 1, 3 and 7 days post-laser exposure (Figure 9 – figure supplement 2). Confocal microscopy revealed GFP-positive processes surrounding PR nuclei, engulfing multiple somata (Figure 9 – figure supplement 2a, c). By 3 days, microglial processes and somas migrated deeper, with processes extending into the distal portions of the PR cells including the inner/outer segments. By two months microglial processes and somas had retreated out of the ONL.

DAPI nuclear stain combined with confocal microscopy not only helped us discern retinal neurons, but also allowed us to differentiate between microglia and phagocytosed PR’s. Two features were different from PR and microglial nuclei. First, each PR displayed a uniform, homogeneous nuclear fluorescence, while microglial nuclei appeared heterogenous and mottled (Figure 9 - figure supplement 2a). Second, microglial nuclei were nearly 3x larger in volume compared to PR nuclei. Microglia had a nuclear volume of 110 ± 42 µm3 and PR nuclei were 35 ± 5 µm3 (p < 0.001, Figure 9 – figure supplement 2b). PR nuclei within microglial phagosomes displayed similar nuclear volume compared to adjacent PR’s in undamaged locations (Figure 9 – figure supplement 2a).

Ex vivo: At lesion sites, neutrophils remain within the retinal vasculature

Corroborating in vivo AOSLO findings we did not find neutrophil aggregation or extravasation in response to laser damage at any of the time points examined. Ex vivo microscopy revealed that neutrophils were found within the vascular network and did not extravasate into the neural retina. Very few neutrophils were found in lesioned retinas, comparable healthy retinas. The few detected neutrophils were remnants of those found within the vascular network at point of death. Detailed 60x Z-stacks revealed “pill-shaped” morphology similar to that seen in vivo, where neutrophils are compressed within capillaries. This morphology contrasts with the spherical shape of neutrophils in response to the EIU positive-control model (Video 2). In one exceptional example (1 day post-damage), we observed two neutrophils in a single z-stack, one in a capillary (1) and one at a capillary branch point (2), but none were observed to have left the confines of the vasculature suggesting they were not recruited by activated microglia (Figure 9c,d, Video 2).

Retinal neutrophil concentrations were quantified from larger z-stacks (796 x 796 µm, Figure 10a). Control locations (n = 2 mice, 4 z-stacks) had 15 ± 8 neutrophils per mm2 of retina whereas lesioned locations (n = 2 mice, 4 z-stacks) had 23 ± 5 neutrophils per mm2 of retina (Figure 10b). The difference between control and lesioned groups were not statistically significant (p = 0.19).

Quantification of neutrophils in laser damaged retinas assessed with ex vivo confocal microscopy over a wide-field.

Earlier version: (A) Representative image displays neutrophils quantified using large-field (796 x 796 µm) z-stacks for control or 1 day after injury time points. In both control and laser injured retinas, neutrophils were sparse and confined to locations within capillaries suggesting they were the native fraction of circulating neutrophils at time of death. Inset displays expanded image of a single neutrophil. Scale bar = 200 µm. (B) Neutrophils quantified and displayed as the number of neutrophils per retinal area. The difference in number of neutrophils in control vs lesioned retinas was not statistically significant (p = 0.19). Error bars display mean + 1 SD.

Revised version: (A) Representative image (maximum intensity projection) displays neutrophils quantified using large-field (796 x 796 µm) z-stacks for control or 1 day after injury time points. In both control and laser injured retinas, neutrophils were sparse and confined to locations within capillaries suggesting they were the native fraction of circulating neutrophils at time of death. Inset displays expanded image of a single neutrophil. Scale bar = 200 µm. (B) Neutrophils quantified and displayed as the number of neutrophils per retinal area. The difference in number of neutrophils in control vs lesioned retinas was not statistically significant (p = 0.19). Error bars display mean + 1 SD.

Finally, when we rendered confocal Z-stacks in cross-sectional view, we discovered that the few fluorescent neutrophils present were found colocalized within the tri-layered vascular stratifications of the mouse retina 32,53,54. None were found in non-vascular layers. Taken together, all ex vivo data indicates that regardless of the robust phagocytic microglial response, neutrophils remain localized within vessels, suggesting that neutrophils do not extravasate or aggregate in response to this laser lesion model. These data corroborate the in vivo findings.

Discussion

Summary

Earlier version: In many retinal diseases, resident microglial populations are found to exhibit cross-talk with systemic immune cells15,18. The complex temporal dynamics between resident and systemic immune cells unfolds from seconds-to-months and has been poorly characterized due to lack of resolution, sufficient contrast and a non-destructive imaging approach that can track changes over time. Here, we overcome these limitations by using phase-contrast and fluorescent imaging with adaptive optics to visualize the interaction of multiple cell types. With this advanced imaging technology, we reveal the absence of neutrophil involvement in response to an acute injury, despite progressive changes in microglial activity and morphology. Here we discuss the implications of such findings in the context of retinal damage and more broadly, to retinal disease.

Revised version: In many retinal diseases, resident microglial populations are found to exhibit cross-talk with systemic immune cells15,18. The complex temporal dynamics between resident and systemic immune cells unfolds from seconds-to-months and has been poorly characterized due to lack of resolution, sufficient contrast and a non-destructive imaging approach that can track changes over time. Here, we overcome these limitations by using phase-contrast and fluorescent imaging with adaptive optics to visualize the interaction of multiple cell types. With this advanced imaging technology, we reveal the absence of neutrophil involvement in response to an acute injury, despite progressive changes in microglial activity and morphology. Here we discuss the implications of such findings in the context of retinal damage and more broadly, retinal disease.

Focal retinal lesions for tracking the immune response

Earlier version: Seminal studies in the brain have used focal light exposure to induce targeted, acute microglial responses.59,60 Since then, it has been a popular damage model. Focal lesions in the retina have been conducted in a number of popular animal models including mouse61,62, macaque63–65 and cats.66,67 Laser damage is also clinically relevant as it is both experienced with accidental laser exposure68,69 as well as purposeful ablation used in photocoagulation and as a therapy for retinal ischemic disease, retinopathy of prematurity70 and diabetic retinopathy71. Thus our results shed light toward the immune response that may be imparted due to phototoxic damage.

Revised version: Seminal studies in the brain have used focal light exposure to induce targeted, acute microglial responses.59,60 Since then, it has been a popular damage model. Focal lesions in the retina have been conducted in a number of popular animal models including mouse61,62, macaque63–65 and cats.66,67 Laser damage is also clinically relevant as it is experienced with accidental laser exposure68,69 as well as purposeful ablation used in photocoagulation and as a therapy for retinal ischemic disease, retinopathy of prematurity70 and diabetic retinopathy71. Thus, our results shed light toward the immune response that may be imparted due to phototoxic damage.

Earlier version: Light damage offers a number of advantages for creating acute damage in the retina. It offers placement precision that is better controlled than other acute and invasive methods, such as incision/poke injuries55, retinal detachment56 and chemical injury57,58. In the mouse, focal light dosage comes with greater axial confinement within the retina by nature of the high numerical aperture of the mouse eye31 which imposes minimal collateral damage to layers above or below the plane of focus (Figure 1a). The damage took on an elliptical form, likely due to: 1) Eye motion from respiration and heart rate which spreads the light over a larger integrative area (rather than line). 2) The impact of focal light scatter. 3) A micron-thin line imparting damage on cells that are many microns across manifesting as an ellipse. The majority of light exposures produced lesions of this elliptical shape. In a few conditions, for the reasons described above, the exposure failed to produce a strong, focal damage phenotype. To improve lesion reproducibility, future experiments should control for subtle eye motion affecting light damage, especially for long exposures.

Revised version: Light damage offers a number of advantages for creating acute damage in the retina. It offers placement precision that is better controlled than other invasive methods, such as incision/poke injuries55, retinal detachment56 and chemical injury57,58. In the mouse, focal light dosage comes with greater axial confinement within the retina by nature of the high numerical aperture of the mouse eye31 which imposes minimal collateral damage to layers above or below the plane of focus (Figure 1a). The damage took on an elliptical form, likely due to: 1) Eye motion from respiration and heart rate which spreads the light over a larger integrative area (rather than line). 2) The impact of focal light scatter. 3) A micron-thin line imparting damage on cells that are many microns across manifesting as an ellipse. Most light exposures produced lesions of this elliptical shape. For the reasons described above, some exposures failed to produce a strong, focal damage phenotype. To improve lesion reproducibility, future experiments should better control for subtle eye motions affecting light distribution, especially for long exposures.

Titration of the 488 nm laser leads to mild retinal damage

Earlier version: While it is beyond the scope of this report to catalog the vast parameter space in which light may impose damage in the retina, it is noteworthy to discuss that the exposure intensity, subtended angle and duration of the light source we used delivered an intentionally mild damage to the retina. Using dosages too high would run the risk of collateral damage to the blood retinal barrier or Bruch’s membrane72 which could cause a choroidal neovascularization phenotype73.

Revised version: While it is beyond the scope of this report to catalog the vast parameter space in which light may impose damage in the retina, it is noteworthy to discuss that the exposure intensity, subtended angle and duration of the light source we used delivered an intentionally mild damage to the retina. Using dosages too high would run the risk of collateral damage to the blood retinal barrier or Bruch’s membrane72 which could cause a choroidal neovascularization phenotype73. This would be undesirable for studying the native/systemic immune response as it could confound the traditional extravasation pathways of the inflammatory system23,74. Instead, we chose our laser exposure condition to impart a weak, but reproducible loss of PRs while minimizing damage to the cells above and below the plane of focus.

Removed in the revised version: This would be undesirable for studying the native/systemic immune response as it could confound the traditional extravasation pathways of the inflammatory system23,74. Instead, we chose our laser exposure condition to impart a weak, but reproducible loss of PRs while minimizing damage to the cells above and below the plane of focus.

Progressive PR loss and morphology of the lesion tracked over time

Earlier version: The amount of light delivered did not immediately ablate PRs as peak cell loss wasn’t observed until 3-7 days after laser exposure, suggesting a progressive damage phenotype of apoptosis or necroptotic death. Histology showed a ∼27% reduction in PR nuclei 7 days post lesion. The total number of PR nuclei within the lesion location returned to values similar to baseline by 2 months, which was surprising given that PRs do not regenerate. Our interpretation is that adjacent PRs spill into the region of loss, similar to what several have described previously63,75. Corroborating the histological finding of mild cell loss, OCT shows no evidence of local edema, cavitation or excavation of the ONL after PR loss (Figure 2a). The return of the appearance, thickness and redistribution of cells within the ablated location indicates that retinal remodeling occurs within 2 months. The concurrence of microglia in these outer retinal bands (where they are normally absent76) support the hypothesis that microglia localized to this region to facilitate a phagocytic injury response and perhaps contribute toward synaptic remodeling11.

Revised version: The amount of light delivered did not immediately ablate PRs as peak cell loss wasn’t observed until 3-7 days after laser exposure, suggesting a progressive damage phenotype of apoptosis or necroptotic death. Histology showed a ∼27% reduction in PR nuclei 7 days post lesion. The total number of PR nuclei within the lesion location returned to values similar to baseline by 2 months, which was surprising given that PRs do not regenerate. Our interpretation is that adjacent PRs spill into the region of loss, similar to what others have described previously63,75. Corroborating the histological finding of mild cell loss, OCT shows no evidence of local edema, cavitation or excavation of the ONL after PR loss (Figure 2a). The return of the appearance, thickness and redistribution of cells within the ablated location indicates that retinal remodeling occurs within 2 months. The concurrence of microglia in these outer retinal bands (where they are normally absent76) support the hypothesis that microglia localized to this region to facilitate a phagocytic injury response and may contribute toward synaptic remodeling11.

CX3CR1-GFP mice exhibit fluorescence not only in microglia

Earlier version: We recognize that the CX3CR1-GFP model can also label systemic cells such as monocytes/macrophages77. While it is possible these cells could infiltrate the retina in response to the lesion, we find it unlikely since there was no indication of the leukocyte extravasation cascade (rolling/crawling/stalled cells) within the nearest retinal vasculature. In addition to microglia, retinal perivascular macrophages and hyalocytes also exhibit GFP fluorescence and thus that these cells may also contribute toward damage resolution.

Revised version: We recognize that the CX3CR1-GFP model can also label systemic cells such as monocytes/macrophages77. While it is possible these cells could infiltrate the retina in response to the lesion, we find it unlikely since there was no indication of the leukocyte extravasation cascade (rolling/crawling/stalled cells) within the nearest retinal vasculature. In addition to microglia, retinal perivascular macrophages and hyalocytes also exhibit GFP fluorescence and these cells may also contribute toward damage resolution.

The hyperreflective phenotype does not arise from microglia or neutrophils

Earlier version: Light damage is known to create hyper-reflective bands in OCT imaging 64,78,79. A common speculation is that the increased backscatter may arise from local inflammatory cells that activate or move into the damage location. In our data, confocal AOSLO and OCT revealed a hyperreflective band at the OPL and ONL after 488 nm light exposure (Figure 2a, b). We found that the hyperreflective bands appeared within 30 minutes after the laser injury, preceding any detectable microglial migration toward the damage location (Figure 2 – figure supplement 1 and Figure 6 – figure supplement 1). We thus conclude that the initial hyperreflective phenotype is not caused by microglial cell activity or aggregation.

Revised version: Light damage is known to create hyper-reflective bands in OCT imaging 64,78,79. A common speculation is that the increased backscatter may arise from local inflammatory cells that activate or move into the damage location. In our data, confocal AOSLO and OCT revealed a hyperreflective band at the OPL/ONL after 488 nm light exposure (Figure 2a, b). We found that the hyperreflective bands appeared within 30 minutes after the laser injury, preceding any detectable microglial migration toward the damage location (Figure 4 – figure supplement 1 and Figure 6 – figure supplement 1). We thus conclude that the initial hyperreflective phenotype is not caused by microglial cell activity or aggregation.

Direct activation of microglia from 488 nm light exposure was minimal

Added in the revised version: It is conceivable that 488 nm light used for either imaging (56 µW) or imparting damage (785 µW) might activate the GFP-containing microglia used here.

Earlier version: It is conceivable that 488 nm light used for either imaging (56 µW) or imparting damage (785 µW) might activate the GFP-containing microglia used here. However, several lines of evidence speak against this possibility. 1) We did not observe photobleaching of CX3CR1 positive cells in response to the damage, suggesting the fluorescence imaging light was insufficient to damage microglia. 2) For retinal injury, the focal plane was adjusted using the AOSLO, such that the dose was axially concentrated onto the outer retina. Thus, the light dosage received by microglia above was defocused and less than the targeted PR layer (Figure 1a). 3) Histology showed no evident necrotic/apoptotic microglial morphologies.

Revised version: However, several lines of evidence speak against this possibility. 1) We did not observe photobleaching of CX3CR1 positive cells in response to the damage, suggesting the light was insufficient to damage microglia. 2) For retinal injury, the focal plane was adjusted such that the dose was axially concentrated onto the outer retina. Thus, the light dosage received by microglia above was defocused and less than the targeted PR layer (Figure 1a). 3) Histology showed no evident necrotic/apoptotic microglial morphologies.

Earlier version: 488 nm laser lesion does not photocoagulate or alter retinal circulation

Revised version: nm laser lesion does not photocoagulate or alter retinal circulation

Earlier version: Despite imparting damage to the PRs, the damage regime used here did not alter the perfusion of the retinal circulation. We show three independent measures that blood flow is uninterrupted, despite PR loss and activation of microglia. 1) Fluorescein angiography (Figure 1b, bottom) revealed an absence of vascular leakage. 2) AOSLO motion contrast vascular maps 24,27,32 displayed persistent blood perfusion inside vessels near laser lesion sites (Figure 4). 3)

Revised version: Despite imparting damage to the PRs, the damage regime used here did not alter the perfusion of the retinal circulation. We show three independent measures that blood flow is uninterrupted, despite PR loss and activation of microglia. 1) Fluorescein angiography (Figure 1b, bottom) revealed an absence of vascular leakage. 2) AOSLO motion contrast vascular maps 24,27,32 displayed persistent blood perfusion inside vessels near lesion sites (Figure 4). 3) Capillary line scans indicate that RBC flux was not modified at lesion locations and fell within the normal range27,48 (Figure 4 – figure supplement 1). Altogether, these three lines of evidence indicate that the lesion did not compromise the blood-retinal barrier or impart perfusion changes within the retinal vasculature.

Removed in the revised version: Capillary line scans indicate that RBC flux was not modified at lesion locations and fell within the normal range27,48 (Figure 4 – figure supplement 1). Altogether, these three lines of evidence indicate that the lesion did not compromise the blood-retinal barrier or impart perfusion changes within the retinal vasculature.

Resident microglia do not need systemic neutrophils for resolution of mild laser-damage

Earlier version: The CNS and retina, unlike other peripheral tissues, cannot suffer from excess inflammation as there may be dire functional consequences. Therefore it is possible that microglia protect against exorbitant inflammation by modulating the recruitment of systemic inflammatory cells.4,15–17 Of these, neutrophils are often one of the first systemic responders. Despite their helpful roles in other tissues, neutrophils can secrete neurotoxic compounds that could present a danger to the CNS.80 Given their conflicted role in the body, we ask the question: to what extent do neutrophils respond to acute neural loss in the retina? Retinal cells are lost with age81 and disease82 and yet, for the organism, visual perception must persist. There are limitations on how generalizable this mild damage to more severe damage or disease phenotypes, but this acute damage model can begin to provide clues about how immune cells interact in response to PR loss. In this laser lesion model, we ablate 27% of the PRs in a 50 µm region.

Revised version: The CNS and retina, unlike other peripheral tissues, cannot suffer from excess inflammation as there may be dire functional consequences. Therefore it is possible that microglia protect against exorbitant inflammation by modulating the recruitment of systemic inflammatory cells.4,15–17 Of these, neutrophils are often one of the first systemic responders. Despite their helpful roles in other tissues, neutrophils can secrete neurotoxic compounds that could present a danger to the CNS.80 Given their conflicted role in the body, we ask the question: to what extent do neutrophils respond to acute neural loss in the retina? Retinal cells are lost with age81 and disease82 and yet, for the organism, visual perception must persist. There are limitations on how generalizable this mild damage is to other damage or disease phenotypes, but this acute damage model can provide clues about how immune cells interact in response to PR loss. In this laser lesion model, we ablate 27% of the PRs in a 50 µm region.

Earlier version: We find that microglia undergo a rapid and progressive response to this injury. We show evidence of PR phagocytosis (Figure 9 - figure supplement 2), interaction with neighboring microglia (Figure 9 – figure supplement 1) and they are also axially positioned to facilitate retinal remodeling (Figure 9). Furthermore, throughout the temporal evolution of the microglial response, we find no evidence of neutrophil recruitment despite the damage being within 10s of microns from retinal vessels that carry them. At the onset of the neutrophil extravasation cascade, endothelial cells in the vicinity of inflamed tissue typically elevate the expression of adhesion molecules, facilitating the adherence and extravasation of circulating neutrophils.83 Furthermore, neutrophils are dependent on priming events as prerequisite to further activation and engagement of their effector functions.84,85 Based on our data, we suggest that although microglia show a strong and lasting activation, at no time point from seconds-to-months, are the damage-associated molecular patterns or chemotactic gradients strong enough to recruit neutrophils in response to this damage. This is evidenced in our data from two key observations: 1) We saw no examples of systemic leukocytes rolling in vessels adjacent to injury locations (Figure 8 – video 1). 2) Using fluorescence, we did not observe adherent or extravasated neutrophils adjacent to imparted PR loss (Figure 8). This may suggest that the region of insult is too small, or that activated microglia are not sufficient for recruiting neutrophils with damage of this magnitude. Perhaps a minimum threshold of neural damage must be met before neutrophils will respond. It is possible that resident microglia facilitate the necessary phagocytic and retinal remodeling response despite release of cytokines from damaged retinal cells that would normally recruit systemic immune cells in peripheral tissues. Such a strategy would benefit the CNS.

Revised version: We find that microglia undergo a rapid and progressive response to this injury. We show evidence of PR phagocytosis (Figure 9 - figure supplement 2), interaction with neighboring microglia (Figure 9 – figure supplement 1) and they are also axially positioned to facilitate retinal remodeling (Figure 9). Furthermore, throughout the temporal evolution of the microglial response, we find no evidence of neutrophil recruitment despite the damage being within 10s of microns from retinal vessels that carry them. At the onset of the neutrophil extravasation cascade, endothelial cells in the vicinity of inflamed tissue typically elevate the expression of adhesion molecules, facilitating the adherence and extravasation of circulating neutrophils.83 Furthermore, neutrophils are dependent on priming events as prerequisite to further activation and engagement of their effector functions.84,85 Based on our data, we suggest that although microglia show a strong and lasting activation, at no time point from seconds-to-months, are the damage-associated molecular patterns or chemotactic gradients strong enough to recruit neutrophils in response to this damage. This is evidenced in our data from two key observations: 1) We saw no examples of systemic leukocytes rolling in vessels adjacent to injury locations (figure 8 – video 1). 2) We did not observe adherent or extravasated neutrophils adjacent to imparted PR loss (Figure 8). This may suggest that the region of insult is too small, or that activated microglia are not sufficient for recruiting neutrophils with damage of this magnitude. Perhaps a minimum threshold of neural damage must be met before neutrophils will respond. It is possible that resident microglia facilitate the necessary phagocytic and retinal remodeling response despite release of cytokines from damaged retinal cells that would normally recruit systemic immune cells in peripheral tissues. Such a strategy would benefit the CNS.

Future work will explore whether there is a threshold magnitude of neuronal cell loss required for recruitment of systemic cells that is unique to the retina. Next studies will examine more severe or widespread injury regimes that provide stronger activating molecular signals and interact with a larger population of systemic cells.

Microglia may inhibit neutrophil activation

Earlier version: Microglia may be involved in a system that protects the CNS from propagating a larger systemic response, potentially exacerbating disease pathologies that would compromise overall CNS function. In another damage model86, they report that tissue-resident macrophages may exhibit the capacity to cloak tissue micro- damage. This offers the possibility that resident immune cells, such as retinal microglia, can handle small insults without inducing a chemokine cascade that may invoke a larger systemic response that could further damage the precious retinal tissue.87 Regardless of the mechanism, we find that despite a robust microglial activation that lasts for weeks, at no time point do they recruit neutrophils. The nuance of this interaction likely represents the fine balance that facilitates a helpful local response within the CNS that does not impart a widespread cytokine storm that may otherwise exacerbate retinal damage.

Revised version: Microglia may be involved in a system that protects the CNS from propagating a larger systemic response, potentially exacerbating disease pathologies that would compromise overall CNS function. In another damage model86, they report that tissue-resident macrophages may exhibit the capacity to cloak tissue micro-damage. This offers the possibility that resident immune cells, such as retinal microglia, can handle small insults without inducing a chemokine cascade that may invoke a larger systemic response that could further damage the precious retinal tissue.87 Regardless of the mechanism, we find that despite a robust microglial activation that lasts for weeks, at no time point do they recruit neutrophils. The nuance of this interaction likely represents the fine balance that facilitates a helpful local response within the CNS that does not impart a widespread cytokine storm that may otherwise exacerbate retinal damage. Further work will explore whether microglia exhibit a cloaking response in the retina, inhibiting neutrophil or other immune cell extravasation/chemotaxis toward lesion sites. We expect such work to be pivotal in understanding the balance that is broken or left unchecked in conditions of autoimmune disease and the umbrella of diseases that comprise the uveitic response, a direct threat to lifelong vision.88

Removed in the revised version: Further work will explore whether microglia exhibit a cloaking response in the retina, inhibiting neutrophil or other immune cell extravasation/chemotaxis toward lesion sites. We expect such work to be pivotal in understanding the balance that is broken or left unchecked in conditions of autoimmune disease and the umbrella of diseases that comprise the uveitic response, a direct threat to lifelong vision.88

Conclusion

Earlier version: Here, we have applied innovative in vivo imaging at the microscopic scale to reveal the cellular immune response to a retina in jeopardy. The dynamic environment of the retina includes a native population of resident microglia and systemic immune cells delivered through the vasculature. These two lines of defense work in concert in the mammalian body and are critical for maintaining retinal homeostasis. In this work we directly study the interaction between microglia and neutrophils, two major classes of immune cells that are implicated in inflammatory initiation, escalation, propagation and debris removal in response to acute geographical injury in the retina. Using cutting-edge retinal imaging modalities, we find that resident microglia become locally activated and regionally responsive to focal laser lesion. Cells migrate away from their stratified locations near plexiform layers of the retina and migrate toward the site of damage within hours to weeks after injury. However, systemic neutrophils, which are typically regarded as first-line responders to tissue damage, are not recruited to this damage despite neutrophils flowing within 10s of microns away from the location of damage (Figure 8 – video 2). Beyond the context of this specific finding, we share this work with the excitement that AOSLO cellular level imaging may reveal the interaction of multiple immune cell types in the living retina. By using fluorophores associated with specific immune cell populations, the complex dynamics that orchestrate the immune response may be examined in this specialized tissue. This work and future studies may reveal further insights to the interactions of single immune cells in the living body in a non-invasive way.

Revised version: Here, we have applied innovative in vivo imaging at the microscopic scale to reveal the cellular immune response to a retina in jeopardy. The dynamic environment of the retina includes a native population of resident microglia and systemic immune cells delivered through the vasculature. These two lines of defense work in concert in the mammalian body and are critical for maintaining retinal homeostasis. In this work we directly study the interaction between microglia and neutrophils, two major classes of immune cells that are implicated in inflammatory initiation, escalation, propagation and debris removal in response to acute geographical injury in the retina. Using cutting-edge retinal imaging modalities, we find that resident microglia become locally activated and regionally responsive to focal laser lesion. They migrate away from their stratified locations near plexiform layers of the retina and toward the site of damage within hours to weeks after injury. However, systemic neutrophils, which are typically regarded as first-line responders to tissue damage, are not recruited to this damage despite neutrophils flowing within 10s of microns away from the location of damage (figure 8 – video 2). Beyond the context of this specific finding, we share this work with the excitement that AOSLO cellular level imaging may reveal the interaction of multiple immune cell types in the living retina. By using fluorophores associated with specific immune cell populations, the complex dynamics that orchestrate the immune response may be examined in this specialized tissue. This work and future studies may reveal further insights to the interactions of single immune cells in the living body in a non-invasive way.

Added in the revised version: Limitations and future directions

Added in the revised version: This work represents some of the earliest reports of single immune cell interactions in the living retina. To narrow the scope, we focus on one type of injury, a targeted elimination of photoreceptors. Thus, these findings are limited to one type of injury that may be experienced in the retina. We expect that these seminal demonstrations will serve as a platform for future studies that examine the large parameter space of retinal damage. These include: 1) Conditions of greater severity with increased power, duration, and extent of light-damage. 2) Models of systemic and local infection. 3) Response to therapy that may modulate the immune response. 4) Examine immune cell activity in models of retinal disease such as diabetic retinopathy, glaucoma, and age-related macular degeneration, each expected to reveal the nuance of the coordinated immune response.